# Factorizations and other fun
Based on work by Andreas Noack

## Reference

https://github.com/JuliaComputing/JuliaBoxTutorials/tree/master/introductory-tutorials/intro-to-julia (github : JuliaComputing/JuliaBoxTutorials/introductory-tutorials/intro-to-julia/)

Topics:

1. [Factorizations](#1)
2. [Special matrix structures](#2)
3. [Generic linear algebra](#3)
4. [Exercises](#4)

### Series

- http://deepstat.tistory.com/45 (01. Getting started)(in English)
- http://deepstat.tistory.com/46 (01. Getting started(한글))
- http://deepstat.tistory.com/47 (02. Strings)(in English)
- http://deepstat.tistory.com/48 (02. Strings(한글))
- http://deepstat.tistory.com/49 (03. Data structures)(in English)
- http://deepstat.tistory.com/50 (03. Data structures(한글))
- http://deepstat.tistory.com/51 (04. Loops)(in English)
- http://deepstat.tistory.com/52 (04. Loops(한글))
- http://deepstat.tistory.com/53 (05. Conditionals)(in English)
- http://deepstat.tistory.com/54 (05. Conditionals(한글))
- http://deepstat.tistory.com/55 (06. Functions)(in English)
- http://deepstat.tistory.com/56 (06. Functions(한글))
- http://deepstat.tistory.com/57 (07. Packages)(in English)
- http://deepstat.tistory.com/58 (07. Packages(한글))
- http://deepstat.tistory.com/59 (08. Plotting)(in English)
- http://deepstat.tistory.com/60 (08. Plotting(한글))
- http://deepstat.tistory.com/61 (09. Julia is fast)(in English)
- http://deepstat.tistory.com/62 (09. Julia is fast(한글))
- http://deepstat.tistory.com/63 (10. Multiple dispatch)(in English)
- http://deepstat.tistory.com/64 (10. Multiple dispatch(한글))
- http://deepstat.tistory.com/65 (11. Basic linear algebra in Julia)(in English)
- http://deepstat.tistory.com/66 (11. Basic linear algebra in Julia(한글))
- http://deepstat.tistory.com/68 (12. Factorizations and other fun(한글))

Before we get started, let's set up a linear system and use `LinearAlgebra` to bring in the factorizations and special matrix structures.

```julia
using LinearAlgebra
A = rand(3, 3)
x = fill(1, (3,))
b = A * x
```

Output:
```
3-element Array{Float64,1}:
 1.4095257746683503
 1.8634312497010534
 1.3670170455218968
```

## Factorizations <a id=1></a>

#### LU factorizations
In Julia we can perform an LU factorization
```julia
PA = LU
``` 
where `P` is a permutation matrix, `L` is lower triangular unit diagonal and `U` is upper triangular, using `lufact`.

Julia allows computing the LU factorization and defines a composite factorization type for storing it.

```julia
Alu = lu(A)
```

Output:
```
LU{Float64,Array{Float64,2}}
L factor:
3×3 Array{Float64,2}:
 1.0       0.0       0.0
 0.615572  1.0       0.0
 0.509803  0.759546  1.0
U factor:
3×3 Array{Float64,2}:
 0.932644   0.799276  0.131511
 0.0       -0.208911  0.428852
 0.0        0.0       0.292487
```

```julia
typeof(Alu)
```

Output:
```
LU{Float64,Array{Float64,2}}
```

The different parts of the factorization can be extracted by accessing their special properties

```julia
Alu.P
```

Output:
```
3×3 Array{Float64,2}:
 0.0  1.0  0.0
 0.0  0.0  1.0
 1.0  0.0  0.0
```

```julia
Alu.L
```

Output:
```
3×3 Array{Float64,2}:
 1.0       0.0       0.0
 0.615572  1.0       0.0
 0.509803  0.759546  1.0
```

```julia
Alu.U
```

Output:
```
3×3 Array{Float64,2}:
 0.932644   0.799276  0.131511
 0.0       -0.208911  0.428852
 0.0        0.0       0.292487
```

Julia can dispatch methods on factorization objects.

For example, we can solve the linear system using either the original matrix or the factorization object.

```julia
A\b
```

Output:
```
3-element Array{Float64,1}:
 1.0000000000000002
 0.9999999999999998
 1.0
```

```julia
Alu\b
```

Output:
```
3-element Array{Float64,1}:
 1.0000000000000002
 0.9999999999999998
 1.0
```

Similarly, we can calculate the determinant of `A` using either `A` or the factorization object

```julia
det(A) ≈ det(Alu)
```

Output:
```
true
```

#### QR factorizations

In Julia we can perform a QR factorization
```
A=QR
``` 

where `Q` is unitary/orthogonal and `R` is upper triangular, using `qrfact`. 

```julia
Aqr = qr(A)
```

Output:
```
LinearAlgebra.QRCompactWY{Float64,Array{Float64,2}}
Q factor:
3×3 LinearAlgebra.QRCompactWYQ{Float64,Array{Float64,2}}:
 -0.398232   0.456045  -0.795886 
 -0.781148  -0.62344    0.0336249
 -0.480853   0.635095   0.604512 
R factor:
3×3 Array{Float64,2}:
 -1.19394  -0.859561  -0.620766
  0.0      -0.205043   0.554298
  0.0       0.0       -0.232786
```

Similarly to the LU factorization, the matrices `Q` and `R` can be extracted from the QR factorization object via

```julia
Aqr.Q
```

Output:
```
3×3 LinearAlgebra.QRCompactWYQ{Float64,Array{Float64,2}}:
 -0.398232   0.456045  -0.795886 
 -0.781148  -0.62344    0.0336249
 -0.480853   0.635095   0.604512
```

```julia
Aqr.R
```

Output:
```
3×3 Array{Float64,2}:
 -1.19394  -0.859561  -0.620766
  0.0      -0.205043   0.554298
  0.0       0.0       -0.232786
```

#### Eigendecompositions

The results from eigendecompositions, singular value decompositions, Hessenberg factorizations, and Schur decompositions are all stored in `Factorization` types.

The eigendecomposition can be computed

```julia
Asym = A + A'
AsymEig = eigen(Asym)
```

Output:
```
Eigen{Float64,Float64,Array{Float64,2},Array{Float64,1}}
eigenvalues:
3-element Array{Float64,1}:
 -0.4684087192598678
  0.9149919921470593
  3.122512789652273 
eigenvectors:
3×3 Array{Float64,2}:
 -0.764548   0.181307  -0.618541
  0.325394  -0.719802  -0.613192
  0.556403   0.670084  -0.491327
```

The values and the vectors can be extracted from the Eigen type by special indexing

```julia
AsymEig.values
```

Output:
```
3-element Array{Float64,1}:
 -0.4684087192598678
  0.9149919921470593
  3.122512789652273
```

```julia
AsymEig.vectors
```

Output:
```
3×3 Array{Float64,2}:
 -0.764548   0.181307  -0.618541
  0.325394  -0.719802  -0.613192
  0.556403   0.670084  -0.491327
```

Once again, when the factorization is stored in a type, we can dispatch on it and write specialized methods that exploit the properties of the factorization, e.g. that A^{-1}=(V\Lambda V^{-1})^{-1}=V\Lambda^{-1}V^{-1}.

```julia
inv(AsymEig)*Asym
```

Output:
```
3×3 Array{Float64,2}:
  1.0           2.77556e-16  -5.9952e-15 
 -1.33227e-15   1.0           4.32987e-15
 -1.65146e-15  -8.04912e-16   1.0
```

## Special matrix structures <a id=2></a>
Matrix structure is very important in linear algebra. To see *how* important it is, let's work with a larger linear system

```julia
n = 1000
A = randn(n,n);
```

Julia can often infer special matrix structure

```julia
Asym = A + A'
issymmetric(Asym)
```

Output:
```
true
```

but sometimes floating point error might get in the way.

```julia
Asym_noisy = copy(Asym)
Asym_noisy[1,2] += 5eps()
```

Output:
```
3.03091928625111
```

```julia
issymmetric(Asym_noisy)
```

Output:
```
false
```

Luckily we can declare structure explicitly with, for example, `Diagonal`, `Triangular`, `Symmetric`, `Hermitian`, `Tridiagonal` and `SymTridiagonal`.

```julia
Asym_explicit = Symmetric(Asym_noisy);
```

Let's compare how long it takes Julia to compute the eigenvalues of `Asym`, `Asym_noisy`, and `Asym_explicit`

```julia
@time eigvals(Asym);
```

Output:
```
0.166887 seconds (112.22 k allocations: 13.368 MiB)
```

```julia
@time eigvals(Asym_noisy);
```

Output:
```
1.073496 seconds (18 allocations: 7.921 MiB)
```

```julia
@time eigvals(Asym_explicit);
```

Output:
```
0.119242 seconds (7.62 k allocations: 8.378 MiB)
```

In this example, using `Symmetric()` on `Asym_noisy` made our calculations about `5x` more efficient :)

#### A big problem
Using the `Tridiagonal` and `SymTridiagonal` types to store tridiagonal matrices makes it possible to work with potentially very large tridiagonal problems. The following problem would not be possible to solve on a laptop if the matrix had to be stored as a (dense) `Matrix` type.

```julia
n = 1_000_000;
A = SymTridiagonal(randn(n), randn(n-1));
@time eigmax(A)
```

Output:
```
0.754640 seconds (484.50 k allocations: 206.799 MiB, 17.43% gc time)
6.370204015035759
```

## Generic linear algebra <a id=3></a>
The usual way of adding support for numerical linear algebra is by wrapping BLAS and LAPACK subroutines. For matrices with elements of `Float32`, `Float64`, `Complex{Float32}` or `Complex{Float64}` this is also what Julia does.

However, Julia also supports generic linear algebra, allowing you to, for example, work with matrices and vectors of rational numbers.

#### Rational numbers
Julia has rational numbers built in. To construct a rational number, use double forward slashes:

```julia
1//2
```

Output:
```
1//2
```

#### Example: Rational linear system of equations
The following example shows how linear system of equations with rational elements can be solved without promoting to floating point element types. Overflow can easily become a problem when working with rational numbers so we use `BigInt`s.

```julia
Arational = Matrix{Rational{BigInt}}(rand(1:10, 3, 3))/10
```

Output:
```
3×3 Array{Rational{BigInt},2}:
 2//5   9//10  2//5 
 2//5   1//2   9//10
 3//10  1//5   1//10
```

```julia
x = fill(1, 3)
b = Arational*x
```

Output:
```
3-element Array{Rational{BigInt},1}:
 17//10
  9//5 
  3//5
```

```julia
Arational\b
```

Output:
```
3-element Array{Rational{BigInt},1}:
 1//1
 1//1
 1//1
```

```julia
lu(Arational)
```

Output:
```
LU{Rational{BigInt},Array{Rational{BigInt},2}}
L factor:
3×3 Array{Rational{BigInt},2}:
 1//1   0//1   0//1
 1//1   1//1   0//1
 3//4  19//16  1//1
U factor:
3×3 Array{Rational{BigInt},2}:
 2//5   9//10     2//5  
 0//1  -2//5      1//2  
 0//1   0//1   -127//160
```

### Exercises <a id=4></a>

#### 11.1
What are the eigenvalues of matrix A?

```
A =
[
 140   97   74  168  131
  97  106   89  131   36
  74   89  152  144   71
 168  131  144   54  142
 131   36   71  142   36
]
```
and assign it a variable `A_eigv`

```julia
using LinearAlgebra
```

```julia
A =[140   97   74  168  131;
    97  106   89  131   36;
    74   89  152  144   71;
    168  131  144   54  142;
    131   36   71  142   36]
```

Output:
```
5×5 Array{Int64,2}:
 140   97   74  168  131
  97  106   89  131   36
  74   89  152  144   71
 168  131  144   54  142
 131   36   71  142   36
```

```julia
A_eigv = eigen(A).values
```

Output:
```
5-element Array{Float64,1}:
 -128.49322764802145 
  -55.887784553057   
   42.752167279318854
   87.16111477514494 
  542.4677301466137
```

<!-- grade_id: cell-4d5f60c8a814c789, points: 0, grade: true, solution: false -->
```julia
@assert A_eigv ==  [-128.49322764802145, -55.887784553056875, 42.7521672793189, 87.16111477514521, 542.4677301466143]
```

Output:
```
Stacktrace:
 [1] top-level scope at In[34]:1
```

#### 11.2 
Create a `Diagonal` matrix from the eigenvalues of `A`.

```julia
A_diag = Diagonal(A_eigv)
```

Output:
```
5×5 Diagonal{Float64,Array{Float64,1}}:
 -128.493     ⋅        ⋅        ⋅         ⋅   
     ⋅     -55.8878    ⋅        ⋅         ⋅   
     ⋅        ⋅      42.7522    ⋅         ⋅   
     ⋅        ⋅        ⋅      87.1611     ⋅   
     ⋅        ⋅        ⋅        ⋅      542.468
```

<!-- grade_id: cell-3b000a3710c9c263, points: 1, grade: true, solution: false -->
```julia
@assert A_diag ==  [-128.493    0.0      0.0      0.0       0.0;
    0.0    -55.8878   0.0      0.0       0.0;
    0.0      0.0     42.7522   0.0       0.0;
    0.0      0.0      0.0     87.1611    0.0;
    0.0 0.0      0.0      0.0     542.468]
```

Output:
```
Stacktrace:
 [1] top-level scope at In[36]:1
```

#### 11.3 
Create a `LowerTriangular` matrix from `A` and store it in `A_lowertri`

```julia
A_lowertri = LowerTriangular(A)
```

Output:
```
5×5 LowerTriangular{Int64,Array{Int64,2}}:
 140    ⋅    ⋅    ⋅   ⋅
  97  106    ⋅    ⋅   ⋅
  74   89  152    ⋅   ⋅
 168  131  144   54   ⋅
 131   36   71  142  36
```

<!-- grade_id: cell-b76cee2b4a8777da, points: 0, grade: true, solution: false -->
```julia
@assert A_lowertri ==  [140    0    0    0   0;
  97  106    0    0   0;
  74   89  152    0   0;
 168  131  144   54   0;
 131   36   71  142  36]
```

### Please let us know how we're doing!
https://tinyurl.com/introJuliaFeedback