# Basic linear algebra in Julia
Author: Andreas Noack Jensen (MIT) (http://www.econ.ku.dk/phdstudent/noack/)
(with edits from Jane Herriman)

## Reference

https://github.com/JuliaComputing/JuliaBoxTutorials/tree/master/introductory-tutorials/intro-to-julia (github : JuliaComputing/JuliaBoxTutorials/introductory-tutorials/intro-to-julia/)

Topics:

1. [Basic operations](#1)
2. [The LinearAlgebra library](#2)
3. [Exercises](#3)

### Series

- http://deepstat.tistory.com/45 (01. Getting started)(in English)
- http://deepstat.tistory.com/46 (01. Getting started(한글))
- http://deepstat.tistory.com/47 (02. Strings)(in English)
- http://deepstat.tistory.com/48 (02. Strings(한글))
- http://deepstat.tistory.com/49 (03. Data structures)(in English)
- http://deepstat.tistory.com/50 (03. Data structures(한글))
- http://deepstat.tistory.com/51 (04. Loops)(in English)
- http://deepstat.tistory.com/52 (04. Loops(한글))
- http://deepstat.tistory.com/53 (05. Conditionals)(in English)
- http://deepstat.tistory.com/54 (05. Conditionals(한글))
- http://deepstat.tistory.com/55 (06. Functions)(in English)
- http://deepstat.tistory.com/56 (06. Functions(한글))
- http://deepstat.tistory.com/57 (07. Packages)(in English)
- http://deepstat.tistory.com/58 (07. Packages(한글))
- http://deepstat.tistory.com/59 (08. Plotting)(in English)
- http://deepstat.tistory.com/60 (08. Plotting(한글))
- http://deepstat.tistory.com/61 (09. Julia is fast)(in English)
- http://deepstat.tistory.com/62 (09. Julia is fast(한글))
- http://deepstat.tistory.com/63 (10. Multiple dispatch)(in English)
- http://deepstat.tistory.com/64 (10. Multiple dispatch(한글))
- http://deepstat.tistory.com/66 (11. Basic linear algebra in Julia(한글))

First let's define a random matrix

```julia
A = rand(1:4,4,4)
```

Output:
```
4×4 Array{Int64,2}:
 4  4  1  1
 1  4  3  3
 4  4  4  4
 3  4  4  2
```

Define a vector of ones

```julia
x = fill(1.0, (4,)) # = fill(1.0, 4)
```

Output:
```
4-element Array{Float64,1}:
 1.0
 1.0
 1.0
 1.0
```

Notice that $A$ has type Array{Int64,2} but $x$ has type Array{Float64,1}. Julia defines the aliases Vector{Type}=Array{Type,1} and Matrix{Type}=Array{Type,2}. 

Many of the basic operations are the same as in other languages <a id=1></a>
#### Multiplication

```julia
b = A*x
```

Output:
```
4-element Array{Float64,1}:
 10.0
 11.0
 16.0
 13.0
```

#### Transposition
As in other languages `A'` is the conjugate transpose, or adjoint

```julia
A'
```

Output:
```
4×4 LinearAlgebra.Adjoint{Int64,Array{Int64,2}}:
 4  1  4  3
 4  4  4  4
 1  3  4  4
 1  3  4  2
```

and we can get the transpose with

```julia
transpose(A)
```

Output:
```
4×4 LinearAlgebra.Transpose{Int64,Array{Int64,2}}:
 4  1  4  3
 4  4  4  4
 1  3  4  4
 1  3  4  2
```

#### Transposed multiplication
Julia allows us to write this without *

```julia
A'A
```

Output:
```
4×4 Array{Int64,2}:
 42  48  35  29
 48  64  48  40
 35  48  42  34
 29  40  34  30
```

#### Solving linear systems 
The problem $Ax=b$ for ***square*** $A$ is solved by the \ function.

```julia
A\b
```

Output:
```
4-element Array{Float64,1}:
 1.0
 1.0
 1.0
 1.0
```

```julia
A^-1 * b
```

Output:
```
4-element Array{Float64,1}:
 0.9999999999999996
 1.0000000000000004
 1.0               
 0.9999999999999991
```

`A\b` gives us the *least squares solution* if we have an overdetermined linear system (a "tall" matrix)

```julia
Atall = rand(4, 2)
```

Output:
```
4×2 Array{Float64,2}:
 0.620125  0.964271
 0.435129  0.392623
 0.352525  0.846371
 0.888477  0.154265
```

```julia
Atall\b
```

Output:
```
2-element Array{Float64,1}:
 12.558191782805045
  7.886509151359657
```

```julia
(Atall'Atall)^-1 * Atall'b
```

Output:
```
2-element Array{Float64,1}:
 12.55819178280504 
  7.886509151359654
```

and the *minimum norm least squares solution* if we have a rank-deficient least squares problem

```julia
v1 = rand(4)
v2 = rand(4)
rankdef = hcat(v1, v1, v2)
```

Output:
```
4×3 Array{Float64,2}:
 0.791797  0.791797  0.631865
 0.132592  0.132592  0.468963
 0.421476  0.421476  0.177085
 0.967094  0.967094  0.202275
```

```julia
rankdef\b
```

Output:
```
3-element Array{Float64,1}:
  5.49389774957956 
  5.493897749579558
 11.26631347260693
```

Julia also gives us the *minimum norm solution* when we have an underdetermined system (a "short" matrix)

```julia
bshort = rand(2)
Ashort = rand(2, 3)
```

Output:
```
2×3 Array{Float64,2}:
 0.913861   0.382365  0.542242
 0.0607539  0.420378  0.750087
```

```julia
Ashort\bshort
```

Output:
```
3-element Array{Float64,1}:
 0.279885521561581  
 0.37741831223794353
 0.6366857785894317
```

# The LinearAlgebra library <a id=2></a>

While much of linear algebra is available in Julia by default (as shown above), there's a standard library named `LinearAlgebra` that brings in many more relevant names and functions. In particular, it provides factorizations and some structured matrix types.  As with all packages, you can bring these additional features into your session with a `using LinearAlgebra`.

### Exercises <a id=3></a>

#### 10.1 
Take the inner product (or "dot" product) of a vector `v` with itself and assign it to variable `dot_v`.



```julia
v = [1,2,3]
```

Output:
```
3-element Array{Int64,1}:
 1
 2
 3
```

```julia
dot_v = v'v
```

Output:
```
14
```

```julia
using LinearAlgebra
```

```julia
dot_v = dot(v,v)
```

Output:
```
14
```

```julia
dot_v = v⋅v # \cdot<tab>
```

Output:
```
14
```

```julia
dot_v = ⋅(v,v) # \cdot<tab>
```

Output:
```
14
```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

#### 10.2 
Take the outer product of a vector v with itself and assign it to variable `cross_v`

```julia
using LinearAlgebra
```

```julia
cross_v = cross(v,v)
```

Output:
```
3-element Array{Int64,1}:
 0
 0
 0
```

```julia
cross_v = v×v # \times<tab>
```

Output:
```
3-element Array{Int64,1}:
 0
 0
 0
```

```julia
cross_v = ×(v,v) # \times<tab>
```

Output:
```
3-element Array{Int64,1}:
 0
 0
 0
```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```